```python
import collections
import string
import itertools
```

---
## Question 1: Infinite sequences (8 pts)
For each of the problems below, write a function which generates the given infinite sequence. We should be able to use your generators to access any entry of the sequence no matter how deep.

**1(a)** (2 pts) The prime numbers are

$$2, 3, 5, 7, 11, 13, 17, \dots$$

Give a generator for the primes. (Note: many algorithms exist for this problem. Yours should be efficient enough that we can use it to generate reasonably large prime numbers.)

<!-- grade_id: primes, points: 0, grade: false, solution: true -->
```python
def primes():
    numbers = itertools.count(2)
    while True:
        prime = next(numbers)
        yield prime
        numbers = itertools.filterfalse(lambda x, p=prime: x % p == 0, numbers)
```

<!-- grade_id: correct_prime, points: 2, grade: true, solution: false -->
```python
"""Check that squares returns the correct output for several inputs"""
p = primes()
plist = [2,3,5,7,11,13,17,19,23,29]
for i in range(10):
    assert next(p) == plist[i]
```

**1(b)** (3 pts) The *ruler sequence* is

$$1, 2, 2, 3, 3, 4, 4, 4, 5, 5, 5, 6, 6, 6, 7, 7, 7, 7\dots$$

The first number (1) appears once; the next two numbers (2 and 3) appear twice, the next three numbers appear three times, etc.

Hint: You are highly encouraged to use relevant functions from itertools to make your solution simple

<!-- grade_id: ruler, points: 0, grade: false, solution: true -->
```python
def ruler():
    number = 1
    for size in itertools.count(1):
        for i in range(size):
            yield from itertools.repeat(number, size)
            number += 1
```

<!-- grade_id: correct_ruler, points: 3, grade: true, solution: false -->
```python
r = ruler()
ruler_numbers = [1, 2, 2, 3, 3, 4, 4, 4, 5, 5, 5, 6, 6, 6, 7, 7, 7, 7, 8, 8, 8, 8]
for n in ruler_numbers:
    assert next(r) == n
```

**1(c)** (3 pts) The look-and-say sequence  

$$1, 11, 21, 1211, 111221, 312211, 13112221, 1113213211, \dots$$

is generated as follows:

1. The first entry is one.
2. The second entry is generated by reading the first entry aloud: "one one"
3. The third entry is generated by reading the second entry aloud: "two ones"
4. The fourth entry is generated by reading the third entry aloud: "one two and one one"
5. The fifth entry is therefore "one one, one two, and two ones"

... and so forth. Note that each entry of the sequence should be a `str` object.

<!-- grade_id: look_say, points: 0, grade: false, solution: true -->
```python
def look_say():
    curr = '1'
    while True:
        yield curr
        next_parts = []
        for ch, group in itertools.groupby(curr):
            count = len(list(group))
            part = str(count) + ch      
            next_parts.append(part)
        curr = ''.join(next_parts)
```

<!-- grade_id: correct_look_say, points: 3, grade: true, solution: false -->
```python
s = look_say()
look_say_first_ten = ['1',
 '11',
 '21',
 '1211',
 '111221',
 '312211',
 '13112221',
 '1113213211',
 '31131211131221',
 '13211311123113112211']
for i in range(10):
    s_i = next(s)
    assert isinstance(s_i, str)
    assert s_i == look_say_first_ten[i]
```

### Question 2: Simple ciphers (5 pts)

A *cipher* is an algorithm for encrypting or decrypting a text message, called *plaintext*, into *ciphertext*. In this exercise, we will assume that all plaintext to be encrypted or decrypted are composed of the lower-case Roman alphabet, without any punctuation or whitespace. Examples of such messages could be

```attackatdawn```

or 

```iamajellyfilleddonut```.

<!-- grade_id: enc_caesar, points: 0, grade: false, solution: true -->
```python
def enc_caesar(s: str, k: int):
    alphabet = string.ascii_lowercase
    res = []
    for ch in s:
        idx = alphabet.index(ch)
        new_idx = (idx + k) % 26
        res.append(alphabet[new_idx])
    return ''.join(res)
```

<!-- grade_id: correct_enc_caesar, points: 2, grade: true, solution: false -->
```python
public_tests = [('thequickbrownfoxjumpsoverthelazydog', 1, 'uifrvjdlcspxogpykvnqtpwfsuifmbazeph'), ('howdoesthiscodehandlewrapping', 18, 'zgovgwklzakugvwzsfvdwojshhafy'),('thekeycanbenegative', -3, 'qebhbvzxkybkbdxqfsb')]
for msg, k, enc in public_tests:
    assert enc_caesar(msg, k) == enc
```

<!-- grade_id: dec_caesar, points: 0, grade: false, solution: true -->
```python
def dec_caesar(s: str, k: int):
    return enc_caesar(s, -k)
```

<!-- grade_id: correct_dec_caesar, points: 3, grade: true, solution: false -->
```python
public_tests = [('thequickbrownfoxjumpsoverthelazydog', 1, 'uifrvjdlcspxogpykvnqtpwfsuifmbazeph'), ('howdoesthiscodehandlewrapping', 18, 'zgovgwklzakugvwzsfvdwojshhafy'),('thekeycanbenegative', -3, 'qebhbvzxkybkbdxqfsb')]
for msg, k, enc in public_tests: 
    assert dec_caesar(enc, k) == msg
```

### Question 3 (5 pts) Valid String.
Given a string s containing just the characters '(', ')', '{', '}', '[' and ']', determine if the input string is valid.

An input string is valid if:

1) Open brackets must be closed by the same type of brackets.
2) Open brackets must be closed in the correct order.
3) Every close bracket has a corresponding open bracket of the same type.


<!-- grade_id: is_valid, points: 0, grade: false, solution: true -->
```python
def is_valid(s: str):
    stack = []
    mapping = {')': '(', '}': '{', ']': '['}
    for ch in s:
        if ch in mapping:
            if not stack or stack.pop() != mapping[ch]:
                return False
        else: 
            stack.append(ch)
    return len(stack) == 0
```

<!-- grade_id: correct_is_valid, points: 5, grade: true, solution: false -->
```python
assert is_valid('()') == True
assert is_valid('([{]}') == False
assert is_valid('([{}])') == True
```